# Lab 5

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

## Lab 5: Recursion

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Double Eights

<hr style="border: 1px solid #fdb515;" />

Write a **recursive** function that takes in a positive integer `n` and determines if its digits contain two adjacent `8`s (that is, two `8`s right next to each other).

> **Hint:** Start by coming up with a recursive plan: the digits of a number have double eights if either ____ (think of something that is straightforward to check) or double eights appear in the rest of the digits.

**Important:** Use recursion! No `for` or `while` loops.

```python
def double_eights(n):
    """Returns whether or not n has two digits in a row that
    are the number 8.

    >>> double_eights(1288)
    True
    >>> double_eights(880)
    True
    >>> double_eights(538835)
    True
    >>> double_eights(284682)
    False
    >>> double_eights(588138)
    True
    >>> double_eights(78)
    False
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("double_eights")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Summation

<hr style="border: 1px solid #fdb515;" />

Write a recursive implementation of `summation`, which takes a positive integer `n` and a function `term`. It applies `term` to every number from `1` to `n` including `n` and returns the sum.

**Important:** Use recursion! No `for` or `while` loops.

```python
def summation(n, term):
    """Return the sum of numbers 1 through n (including n) with term applied to each number.

    >>> summation(5, lambda x: x * x * x) # 1^3 + 2^3 + 3^3 + 4^3 + 5^3
    225
    >>> summation(9, lambda x: x + 1)     # 2 + 3 + 4 + 5 + 6 + 7 + 8 + 9 + 10
    54
    >>> summation(5, lambda x: 2**x)      # 2^1 + 2^2 + 2^3 + 2^4 + 2^5
    62
    """
    assert n >= 1
    "*** YOUR CODE HERE ***"

```

```python
grader.check("summation")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Ten-Pairs

<hr style="border: 1px solid #fdb515;" />

Write a function that takes a positive integer `n` and returns the number of ten-pairs it contains.  A ten-pair is a pair of digits within `n` that sums to 10.

The number 7,823,952 has 3 ten-pairs. The first and fourth digits sum to 7+3=10, the second and third digits sum to 8+2=10, and the second and last digit sum to 8+2=10.

Important notes:
- A digit can be part of more than one ten-pair.
- One 5 does not make a ten-pair with itself.

> *Recommended*: Complete and use the helper function `count_digit` to calculate
how many times a digit appears in `n`.

**Important:** Use recursion! No `for` or `while` loops.

```python
def ten_pairs(n):
    """Return the number of ten-pairs within positive integer n.

    >>> ten_pairs(7823952)
    3
    >>> ten_pairs(55055)
    6
    >>> ten_pairs(9641469)
    6
    """
    "*** YOUR CODE HERE ***"

def count_digit(n, digit):
    """Return how many times digit appears in n.

    >>> count_digit(55055, 5)
    4
    """
    "*** YOUR CODE HERE ***"


```

```python
grader.check("ten_pairs")
```

Tests for double_eights (points: 1)
```
>>> double_eights(1288) == True
True
```
```
>>> double_eights(880) == True
True
```
```
>>> double_eights(538835) == True
True
```
```
>>> double_eights(284682) == False
True
```
```
>>> double_eights(588138) == True
True
```
```
>>> double_eights(78) == False
True
```


Tests for summation (points: 1)
```
>>> summation(5, lambda x: x * x * x) == 225
True
```
```
>>> summation(9, lambda x: x + 1) == 54
True
```
```
>>> summation(5, lambda x: 2 ** x) == 62
True
```


Tests for ten_pairs (points: 1)
```
>>> assert ten_pairs(7823952) == 3
>>> assert ten_pairs(55055) == 6
>>> assert ten_pairs(9641469) == 6
>>> assert count_digit(55055, 5) == 4
```
