# Lab 9

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

## Lab 9: Trees

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Efficiency

<hr style="border: 1px solid #fdb515;" />

Recall that the order of growth of a function expresses how long it
takes for the function to run, and is defined in terms of the
function's input sizes.

For example, let's say that we have the function `get_x` which is
defined as follows:

```python
    def get_x(x):
        return x
```

`get_x` has one expression in it. That one expression takes the same
amount of time to run, no matter what x is, or more importantly, how
large x gets. This is called constant time, or $O(1)$.

The main two ways that a function in your program will get a running
time different than just constant time is through either iteration or
recursion. Let's start with some iteration examples!

The (simple) way you figure out the running time of a particular while loop is to simply count the cost of each operation in the body of the while loop, and then multiply that cost by the number of times that the loop runs. For example, look at the following method with a loop in it:

```python
    def foo(n):
        i = 1
        sum = 0
        while i <= n:
            sum = sum + i
            i = i + 1
        return sum
```

This loop has two statements in it `sum = sum + i` and `i = i + 1.`
Each statement is a constant time operation, since the amount of time each statement takes does not depend on the input to the function `n`. In C88Cx, we are not concerned with how long primitive functions such as addition, multiplication, and variable assignment take to run. Rather we are concerned with _how many more times a loop is executed_ or _how many more recursive calls_ occur as the input increases. In this example, we execute the loop `n` times, and
for each iteration, we only execute constant time operations, so we get an order of growth of $O(n$).

Here are a couple of basic functions, along with their running times.
Try to understand why they have the given running time.

---
**$O(n)$**

```python
    def bar(n):
        i = 1
        a = 1
        b = 0
        while i <= n:
            temp = a
            a = a + b
            b = temp
            i = i + 1
        return a
```

--- 

**$O(n^2)$**
```python
    def bar(n):
        sum = 0
        a, b = 0, 0
        while a < n:
            while b < n:
                sum += (a*b)
                b += 1
            b = 0
            a += 1
        return sum
```

**_Note_: There is nothing to submit for this part.** But doing these problems will be good practice. The solutions are given right below the question and to see them you must highlight them.

---

## Baz


What is the asymptotic run time of the baz function?

```python
    def baz(n):
        i = 1
        sum = 0
        while i <= n:
            sum += bam(i)
            i += 1
        return sum

    def bam(n):
        i = 1
        sum = 0
        while i <= n:
            sum += i
            i += 1
        return sum
```

Highlight the text on the line below this line to see the solution:

Answer: <span style="color:white">O(_n_<sup>2</sup>))</span>.


---

## Bonk

What is the asymptotic run time of the bonk function?

```python
def bonk(n):
    sum = 0
    while n >= 2:
        sum += n
        n = n / 2
    return sum
```

Highlight the text on the line below this line to see the solution:
Answer: <span style="color:white">O(log(_n_))</span>.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Trees

<hr style="border: 1px solid #fdb515;" />

Our `Tree` consists of an `value` and a list of its
`branches`. To create a tree and access its root and branches, use the
following constructor and selectors:

* Constructor
    * `Tree(value, branches)`: creates a tree object with the given `root` and list of `branches`. 
    * `Tree(value)`: creates a tree object with no branches to make a leaf, similar to what we saw with linked lists.
* Selectors
    * `tree.value`: returns the value of the root of the `tree`.
    * `tree.branches`: returns the list of branches of the given `tree`.
    * `tree.is_leaf()`: returns `True` if `tree`'s list of `branches` is empty, and `False` otherwise. Notice this is a function call

For example, the tree generated by

```python
    t = Tree(1, [Tree(2),
                 Tree(3, [Tree(4), Tree(5)]),
                 Tree(6, [Tree(7)])])
```

would look like this:

```markdown
       1
     / | \
    2  3  6
      / \  \
     4   5  7
```

It may be easier to visualize this translation by formatting the code
like this:

```python
    t = Tree(
        1,
        [
            Tree(2),
            Tree(3, [Tree(4), Tree(5)] ),
            Tree(6, [Tree(7)] )
        ]
    )
```

To extract the number `3` from this tree, which is the value of the root of its second branch, we would do this:
```python
    t.branches[1].value
```

Here is the implementation of the tree class, with the `__repr__` function giving you what you need to type to create a tree when you enter an instance of the tree into the interpreter, and the `__str__` function giving you a pretty version of a tree when you print it.

```python
class Tree:
    def __init__(self, value, branches=()):
        self.value = value
        for branch in branches:
            assert isinstance(branch, Tree)
        self.branches = list(branches)

    def __repr__(self):
        if self.branches:
            branches_str = ', ' + repr(self.branches)
        else:
            branches_str = ''
        return 'Tree({0}{1})'.format(self.value, branches_str)

    def __str__(self):
        def print_tree(t, indent=0):
            tree_str = '  ' * indent + str(t.value) + "\n"
            for b in t.branches:
                tree_str += print_tree(b, indent + 1)
            return tree_str
        return print_tree(self).rstrip()

    def is_leaf(self):
        return not self.branches
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Search

<hr style="border: 1px solid #fdb515;" />

Write a function `search` that returns the `Tree`, whose root node is the given value if it exists and None if it does not. You can assume all values are unique.

```python
def search(t, value):
    """Searches for and returns the Tree whose value is equal to value if
    it exists and None if it does not. Assume unique values.

    >>> t = Tree(1, [Tree(3, [Tree(5)]), Tree(7)])
    >>> search(t, 10)
    >>> search(t, 5)
    Tree(5)
    >>> search(t, 1)
    Tree(1, [Tree(3, [Tree(5)]), Tree(7)])
    >>> search(t, 7)
    Tree(7)
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("search")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Square

<hr style="border: 1px solid #fdb515;" />

Write a function `label_squarer` that mutates a `Tree` with numerical labels so
that each label is squared.

```python
def label_squarer(t):
    """Mutates a Tree t by squaring all its elements.

    >>> t = Tree(1, [Tree(3, [Tree(5)]), Tree(7)])
    >>> label_squarer(t)
    >>> t
    Tree(1, [Tree(9, [Tree(25)]), Tree(49)])
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("square_tree")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Add Leaves

<hr style="border: 1px solid #fdb515;" />


Implement `add_d_leaves`, a function that takes in a `Tree` instance `t` and a number `v`.

We define the depth of a node in `t` to be the number of edges from the root to that node. The depth of root is therefore 0.

For each node in the tree, you should add `d` leaves to it, where `d` is the depth of the node. Every added leaf should have a label of `v`. If the node at this depth has existing branches, you should add these leaves to the end of that list of branches.

For example, you should be adding 1 leaf with label `v` to each node at depth 1, 2 leaves to each node at depth 2, and so on.

Here is an example of a tree `t` (shown on the left) and the result after `add_d_leaves` is applied with `v` as 5.

![add](add_leaves.png)

> **Hint:** Use a helper function to keep track of the depth!

> **Hint:** Be careful about when you add the new leaves to the tree.
> We only want to add leaves to the original nodes in the tree, not to
> the nodes we added.

```python
def add_d_leaves(t, v):
    """Add d leaves containing v to each node at every depth d.

    >>> t_one_to_four = Tree(1, [Tree(2), Tree(3, [Tree(4)])])
    >>> print(t_one_to_four)
    1
      2
      3
        4
    >>> add_d_leaves(t_one_to_four, 5)
    >>> print(t_one_to_four)
    1
      2
        5
      3
        4
          5
          5
        5

    >>> t0 = Tree(9)
    >>> add_d_leaves(t0, 4)
    >>> t0
    Tree(9)
    >>> t1 = Tree(1, [Tree(3)])
    >>> add_d_leaves(t1, 4)
    >>> t1
    Tree(1, [Tree(3, [Tree(4)])])
    >>> t2 = Tree(2, [Tree(5), Tree(6)])
    >>> t3 = Tree(3, [t1, Tree(0), t2])
    >>> print(t3)
    3
      1
        3
          4
      0
      2
        5
        6
    >>> add_d_leaves(t3, 10)
    >>> print(t3)
    3
      1
        3
          4
            10
            10
            10
          10
          10
        10
      0
        10
      2
        5
          10
          10
        6
          10
          10
        10
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("add_d_leaves")
```

Tests for add_d_leaves (points: 1)
```
>>> t_one_to_four = Tree(1, [Tree(2), Tree(3, [Tree(4)])])
>>> add_d_leaves(t_one_to_four, 5)
>>> print(t_one_to_four)
1
  2
    5
  3
    4
      5
      5
    5
```
```
>>> t0 = Tree(9)
>>> add_d_leaves(t0, 4)
>>> repr(t0) == 'Tree(9)'
True
```
```
>>> t1 = Tree(1, [Tree(3)])
>>> add_d_leaves(t1, 4)
>>> repr(t1) == 'Tree(1, [Tree(3, [Tree(4)])])'
True
```
```
>>> t1 = Tree(1, [Tree(3, [Tree(4)])])
>>> t2 = Tree(2, [Tree(5), Tree(6)])
>>> t3 = Tree(3, [t1, Tree(0), t2])
>>> add_d_leaves(t3, 10)
>>> print(t3)
3
  1
    3
      4
        10
        10
        10
      10
      10
    10
  0
    10
  2
    5
      10
      10
    6
      10
      10
    10
```


Tests for search (points: 1)
```
>>> t = Tree(1, [Tree(3, [Tree(5)]), Tree(7)])
>>> assert search(t, 10) is None
>>> assert search(t, 5).value == 5
>>> assert search(t, 1).value == 1
>>> assert len(search(t, 1).branches) == 2
>>> assert search(t, 7).value == 7
```


Tests for square_tree (points: 1)
```
>>> import io
>>> from contextlib import redirect_stdout
>>> t = Tree(1, [Tree(3, [Tree(5)]), Tree(7)])
>>> label_squarer(t)
>>> f = io.StringIO()
>>> with redirect_stdout(f):
...     print(t)
>>> output = f.getvalue().split('\n')
>>> output.remove('')
>>> output == ['1', '  9', '    25', '  49']
True
```
