# Lab 4

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

## Lab 4: Abstract Data Types

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Distance

<hr style="border: 1px solid #fdb515;" />

We will now use those selectors to write `distance`, which computes the distance between two city objects. Recall that the distance between two coordinate pairs, `(x1, y1)` and `(x2, y2)` can be found by calculating the `sqrt` of `(x1 - x2)**2 + (y1 - y2)**2`. We have already imported `sqrt` for your convenience, so use that and the appropriate selectors to fill in the function.

```python
# Treat all the following code as being behind an abstraction layer, you shouldn't need to look at it!

def make_city(name, lat, lon):
    """
    >>> city = make_city('Berkeley', 0, 1)
    >>> get_name(city)
    'Berkeley'
    >>> get_lat(city)
    0
    >>> get_lon(city)
    1
    """
    if change_abstraction.changed:
        return {"name" : name, "lat" : lat, "lon" : lon}
    else:
        return [name, lat, lon]

def get_name(city):
    """
    >>> city = make_city('Berkeley', 0, 1)
    >>> get_name(city)
    'Berkeley'
    """
    if change_abstraction.changed:
        return city["name"]
    else:
        return city[0]

def get_lat(city):
    """
    >>> city = make_city('Berkeley', 0, 1)
    >>> get_lat(city)
    0
    """
    if change_abstraction.changed:
        return city["lat"]
    else:
        return city[1]

def get_lon(city):
    """
    >>> city = make_city('Berkeley', 0, 1)
    >>> get_lon(city)
    1
    """
    if change_abstraction.changed:
        return city["lon"]
    else:
        return city[2]

def change_abstraction(change):
    change_abstraction.changed = change

change_abstraction.changed = False

from math import sqrt

def distance(city1, city2):
    """
    >>> city1 = make_city('city1', 0, 1)
    >>> city2 = make_city('city2', 0, 2)
    >>> distance(city1, city2)
    1.0
    >>> city3 = make_city('city3', 6.5, 12)
    >>> city4 = make_city('city4', 2.5, 15)
    >>> distance(city3, city4)
    5.0
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("distance")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Close city

<hr style="border: 1px solid #fdb515;" />

Implement `closer_city`, a function that takes a latitude,
longitude, and two cities, and returns the name of the city that is
relatively closer to the provided latitude and longitude.

You may only use selectors and constructors (introduced above) for
this question. You may also use the `distance` function
defined above. Remember, the point of data abstraction,
as we've said, is that we do not need to know how an abstract data type is
implemented, but rather just how we can interact with and use the data type.

```python
def closer_city(lat, lon, city1, city2):
    """Returns the name of either city1 or city2, whichever is closest
    to coordinate (lat, lon).

    >>> berkeley = make_city('Berkeley', 37.87, 112.26)
    >>> stanford = make_city('Stanford', 34.05, 118.25)
    >>> closer_city(38.33, 121.44, berkeley, stanford)
    'Stanford'
    >>> bucharest = make_city('Bucharest', 44.43, 26.10)
    >>> vienna = make_city('Vienna', 48.20, 16.37)
    >>> closer_city(41.29, 174.78, bucharest, vienna)
    'Bucharest'
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("city")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Closer City Abstraction

<hr style="border: 1px solid #fdb515;" />

Run the following tests to make sure that you are using abstraction barriers correctly! You should not need to change your code from the previous question to pass this test.

```python
def check_abstraction():
    """
    There's nothing for you to do for this function, it's just here for the extra doctest
    >>> change_abstraction(True)
    >>> city1 = make_city('city1', 0, 1)
    >>> city2 = make_city('city2', 0, 2)
    >>> distance(city1, city2)
    1.0
    >>> city3 = make_city('city3', 6.5, 12)
    >>> city4 = make_city('city4', 2.5, 15)
    >>> distance(city3, city4)
    5.0
    >>> berkeley = make_city('Berkeley', 37.87, 112.26)
    >>> stanford = make_city('Stanford', 34.05, 118.25)
    >>> closer_city(38.33, 121.44, berkeley, stanford)
    'Stanford'
    >>> bucharest = make_city('Bucharest', 44.43, 26.10)
    >>> vienna = make_city('Vienna', 48.20, 16.37)
    >>> closer_city(41.29, 174.78, bucharest, vienna)
    'Bucharest'
    >>> change_abstraction(False)
    """
```

```python
grader.check("city_abstraction")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Menu

<hr style="border: 1px solid #fdb515;" />

For this question, let's combine ADTs and dictionaries to make an ADT that represents a food object. The `menu` ADT has a constructor `make_food`. There are three selector functions: `get_name`, `get_size`, and `get_price`. When you implement an ADT, you have the freedom to choose how you want to represent it. Some common data structures used to represent ADTs are lists and dictionaries. The `city` ADT was implemented using a list. Now, we will use dictionaries to implemt our `menu` ADT. You must use a dictionary for this question, or else you will not pass the tests.

```python
def make_food(name, size, price):
    """
    >>> fries = make_food('Fries', 'Small', 3)
    >>> isinstance(fries, dict)
    True
    """
    "*** YOUR CODE HERE ***"

def get_food_name(food):
    """
    >>> fries = make_food('Fries', 'Small', 3)
    >>> get_food_name(fries)
    'Fries'
    >>> pizza = make_food('Pizza', 'Medium', 18)
    >>> get_food_name(pizza)
    'Pizza'
    """
    "*** YOUR CODE HERE ***"

def get_size(food):
    """
    >>> fries = make_food('Fries', 'Small', 3)
    >>> get_size(fries)
    'Small'
    >>> pizza = make_food('Pizza', 'Medium', 18)
    >>> get_size(pizza)
    'Medium'
    """
    "*** YOUR CODE HERE ***"

def get_price(food):
    """
    >>> fries = make_food('Fries', 'Small', 3)
    >>> get_price(fries)
    3
    >>> pizza = make_food('Pizza', 'Medium', 18)
    >>> get_price(pizza)
    18
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("make_food")
```

Tests for city (points: 1)
```
>>> berkeley = make_city('Berkeley', 37.87, 112.26)
>>> stanford = make_city('Stanford', 34.05, 118.25)
>>> assert closer_city(38.33, 121.44, berkeley, stanford) == 'Stanford'
```
```
>>> bucharest = make_city('Bucharest', 44.43, 26.1)
>>> vienna = make_city('Vienna', 48.2, 16.37)
>>> assert closer_city(41.29, 174.78, bucharest, vienna) == 'Bucharest'
```


Tests for city_abstraction (points: 1)
```
>>> change_abstraction(True)
>>> city1 = make_city('city1', 0, 1)
>>> city2 = make_city('city2', 0, 2)
>>> assert distance(city1, city2) == 1.0
>>> city3 = make_city('city3', 6.5, 12)
>>> city4 = make_city('city4', 2.5, 15)
>>> assert distance(city3, city4) == 5.0
>>> berkeley = make_city('Berkeley', 37.87, 112.26)
>>> stanford = make_city('Stanford', 34.05, 118.25)
>>> assert closer_city(38.33, 121.44, berkeley, stanford) == 'Stanford'
>>> bucharest = make_city('Bucharest', 44.43, 26.1)
>>> vienna = make_city('Vienna', 48.2, 16.37)
>>> assert closer_city(41.29, 174.78, bucharest, vienna) == 'Bucharest'
>>> change_abstraction(False)
```


Tests for distance (points: 1)
```
>>> city1 = make_city('city1', 0, 1)
>>> city2 = make_city('city2', 0, 2)
>>> assert distance(city1, city2) == 1.0
```


Tests for make_food (points: 1)
```
>>> fries = make_food('Fries', 'Small', 3)
>>> pizza = make_food('Pizza', 'Medium', 18)
>>> assert isinstance(fries, dict)
>>> assert get_food_name(fries) == 'Fries'
>>> assert get_food_name(pizza) == 'Pizza'
>>> assert get_size(fries) == 'Small'
>>> assert get_size(pizza) == 'Medium'
>>> assert get_price(fries) == 3
>>> assert get_price(pizza) == 18
```
