# Lab 0

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

## Lab 0: Getting Started

Welcome to the first lab of Data 88Cx!

The goals of this lab notebook:

* Access all of the technology we'll be using in this course;
* Familiarize yourself with the Jupyter Notebook environment;
* Use arithmetic operators in Python; and
* Work with Python names.

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab00.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Introducing Jupyter Notebooks

## Cells

**Not these!** &rarr;

<div align=center>
    
<img src='cell_graphic.jpeg' alt="Depiction of biological cell" width=200>
   
</div>

<br>

Jupyter Notebooks are made up of cells. There are two types of cells in a Jupyter Notebook: code cells and [Markdown](https://www.markdownguide.org/getting-started/) cells. 
- Code cells are where we write our Python code.
- Markdown cells allow us to write text, like the text you're reading right now.

## Code cells

Running a code cell will execute all of the code it contains.  Any output produced by the code will appear below the cell. Notice the `In:` and `Out:` statements to the left of the cells, followed by a number. That number designates the order in which the cell was run. For example, the number `[3]` next to a cell would indicate that this cell is the third cell that has been run in the notebook so far, or that this is the third time you've run this cell if you've re-run the cell multiple times.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle. Next, do one of the following:
* **Keyboard shortcut**: Hold down the `Shift` key and press `Return` or `Enter`.
* **Button click**: Click the "Run" button in the cell toolbar above.

While the above two options give no difference in computer runtime performance, it's useful to internalize some common keyboard shortcuts early in your Jupyter experience (think of it like learning "Ctrl-C" or "Command-C" to copy instead of "Right click -> Copy to Clipboard"). While they're not required for this course, keyboard shortuts will greatly improve your quality of life when programming Python (just trust us on this).


We can use code cells to do arithmetic in Python. Try running the code cells below:

```python
3 + 4 ** 2
```

The order of evaluation is what you might expect: parentheses, then exponentiation, then multiplication and division, then addition and subtraction; from left to right. 

```python
2 *(3 + 10)
```

You can use parentheses to modify the typical order of evaluation. **How would the above output change, if you removed the parentheses?**

```python
4 ** 2
1 + 2 + 4 + 8 + 16 + 32 + 64
```

**Note that only the value from evaluating the last expression is displayed;** the output from `4 ** 2` isn't shown at all. This is expected behavior.

```python
10000
```

Also notice that this cell contains a very simple expression that evaluates to `1000`.

```python
10 + 20 + 30 - max(10, 20, 30) - min(10, 20, 30)
```

The output you see from running these five code cells are the values `19`, `16`, `127`, `1000` and `20` displayed, respectively. 

Comments in code cells are denoted using a hashtag (`#`). Comments are used to explain what code does when necessary (or, occasionally in our case, to provide instructions).

```python
# This is a comment!
19 - 18
```

If you try and use syntax Python doesn't know about, it will give you an error.

```python
# 3 + * 4 
```

```python
# Why does this work while the above cell causes an error?
3 * - 4
```

<br/>

---

## Markdown Cells


To edit an existing Markdown cell, double click it and change whatever you want. To *render* (display) it, you can "run" it like you'd run a code cell.

**Edit the following cell to contain your favorite color.**

*Note:* This is how questions will be formatted in your homework assignments.

My favorite color is maroon. *(Type your answer here, replacing the text)*

---

## Converting cells

To convert a cell from code to Markdown (or vice versa), you should use the dropdown menu in the cell toolbar. For instance, to convert the following cell from Markdown to code, you should:
1. Click the cell.
2. Click "Markdown" in the cell toolbar, and change it to "Code."

Your cell toolbar should then look like this:

<img src="code-celltoolbar.png" alt="depiction of cell toolbar" width="50%"/>

**Change the following cell from a Markdown cell to a code cell**, and run it by pressing the "Run" button or using the keyboard shortcut `Shift + Enter`.

3 + 4 + 5  # Change this to a code cell

---

## Adding and deleting cells

To add cells, you can go to "Insert" in the toolbar at the top, then click "Insert cell Above" or "Insert cell Below." To delete cells, you can go to "Edit" in the toolbar at the top and click "Delete cells". (Be very careful when deleting a cell.)

For your convenience, we provide this screenshot that details how to switch between "Edit mode" and "Command mode", which is relevant if you want to learn keyboard shortcuts. Again, keyboard shortcuts are not required for this course, but learning some now will greatly improve your programming quality of life.

<img src='edit-mode.png' alt='Jupyter Notebook Edit mode vs Command mode visual guide (cell border indicators).' width=400>

Some useful keyboard shortcuts:

| Action | Mode | Keyboard shortcut |
| :--- | :--- | :---
| Run cell + jump to next cell | Either (puts you in edit mode) | `SHIFT + ENTER` |
| Save notebook | Either | `CTRL/CMD + S` |
| Switch to command mode | Either (puts you in command mode) | `ESCAPE` |
| Switch to edit mode | Command | `ENTER` | 
| Comment out the current line | Edit | `CTRL/CMD + /` |
| Create new cell above/below | Command |  `A/B` |
| Delete cell | Command | `DD` |
| Convert cell to Markdown | Command | `M` |
| Convert cell to code | Command | `Y` |
| Show all shortcuts | Command |  `H` |


## Insert Cell

**Add a Markdown cell below this one, and in it write your middle name, or write "I don't have a middle name".**

```python
# The cell you create should be above this one!
```

---

## Delete Cell

Delete the cell below.

```python
# This should be deleted!
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## What Would Python Do? (WWPD)

<hr style="border: 1px solid #fdb515;" />

One component of lab assignments is to predict how the Python will behave. For the following questions, assign the variable to the **output** of various statements/expressions.

**Note:** For the following questions that ask you to write code, it will usually be followed with a cell that says something like `otter.grade('q1')` that you'll need to run in order for your work to be graded. 

---

## a)

```python
>>> x = 20
>>> x + 2
```

```python
one_a = ...
```

```python
grader.check("wwpd_a")
```

---

## b)

```python
>>> x
```

```python
one_b = ...
```

```python
grader.check("wwpd_b")
```

---

## c) 

```python
>>> y = 5
>>> y = y + 3
>>> y * 2
```

```python
one_c = ...
```

```python
grader.check("wwpd_c")
```

---

## d)

```python
>>> y + x
```

```python
one_d = ...
```

```python
grader.check("wwpd_d")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Implementing the 1868 Function

<hr style="border: 1px solid #fdb515;" />

Labs will often ask you to implement functions. You should see a function called `eighteen_sixty_eight` below that has a blank
`return` statement. That blank is the only part you should change. Come up with the most creative expression that evaluates to 1868 using only numbers and the +, *, and - operators (or ** and % if you'd like)

**Fun Fact**: The [University of California](https://www.berkeley.edu/about/) was founded in 1868, born out of a vision in the State Constitution of a university that would “contribute even more than California’s gold to the glory and happiness of advancing generations.”

```python
def eighteen_sixty_eight():
    return ______

eighteen_sixty_eight() # This line should output 1868 after being run
```

```python
grader.check("eighteen_sixty_eight")
```

Tests for eighteen_sixty_eight (points: 1)
```
>>> eighteen_sixty_eight()
1868
```


Tests for wwpd_a (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_a).encode()).hexdigest()
'785f3ec7eb32f30b90cd0fcf3657d388b5ff4297f2f9716ff66e9b69c05ddd09'
```


Tests for wwpd_b (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_b).encode()).hexdigest()
'f5ca38f748a1d6eaf726b8a42fb575c3c71f1864a8143301782de13da2d9202b'
```


Tests for wwpd_c (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_c).encode()).hexdigest()
'b17ef6d19c7a5b1ee83b907c595526dcb1eb06db8227d650d5dda0a9f4ce8cd9'
```


Tests for wwpd_d (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_d).encode()).hexdigest()
'59e19706d51d39f66711c2653cd7eb1291c94d9b55eb14bda74ce4dc636d015a'
```
