# Lab 8

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

## Lab 8: Inheritance + Linked Lists

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Inheritance

<hr style="border: 1px solid #fdb515;" />

## Checking Account

Consider the `Account` class provided below, which represents the attributes and behaviors of a bank account.

```python
class Account:
    """A bank account that allows deposits and withdrawals.

    >>> eric_account = Account('Eric')
    >>> eric_account.deposit(1000000)   # depositing my paycheck for the week
    1000000
    >>> eric_account.transactions
    [('deposit', 1000000)]
    >>> eric_account.withdraw(100)      # buying dinner
    999900
    >>> eric_account.transactions
    [('deposit', 1000000), ('withdraw', 100)]
    """

    interest = 0.02

    def __init__(self, account_holder):
        self.balance = 0
        self.holder = account_holder
        self.transactions = []

    def deposit(self, amount):
        """Increase the account balance by amount and return the
        new balance.
        """
        self.transactions.append(('deposit', amount))
        self.balance = self.balance + amount
        return self.balance

    def withdraw(self, amount):
        """Decrease the account balance by amount and return the
        new balance.
        """
        self.transactions.append(('withdraw', amount))
        if amount > self.balance:
            return 'Insufficient funds'
        self.balance = self.balance - amount
        return self.balance
```

We want to create a new class called `CheckingAccount` which inherits from `Account`. We'd like to be able to cash checks to our `CheckingAccount`. To do this, you will have to implement the following:

1. Implement the `Check` class's `__init__` and `mark_deposited` methods, which represents a check that can be deposited into an account. Read the doctests in the `Check` class to learn about the behavior of these methods. Hint: Think about what sort of attributes our `Check` class should have.
2. Implement the `CheckingAccount` class's `deposit_check` method. It will take a `Check` object as an argument.
    - If the check was already deposited, print `The police have been notified`
    - If the check is NOT payable to the checking account holder, print `The police have been notified`
    - Otherwise, mark the check as deposited and deposit the amount into the checking account

The doctests for this problem are extremely important and valuable in understanding how this code works and how the `CheckingAccount` and `Check` classes should interact with each other.

```python
class Check:
    """A Check object that can be deposited.

    >>> eric_check = Check("Eric", 30)  # 30 dollars, payable to Eric
    >>> eric_check.payable_to
    'Eric'
    >>> eric_check.amount
    30
    >>> eric_check.deposited
    False
    >>> eric_check.mark_deposited()
    >>> eric_check.deposited
    True
    """
    def __init__(self, payable_to, amount):
        "*** YOUR CODE HERE ***"

    def mark_deposited(self):
        "*** YOUR CODE HERE ***"

class CheckingAccount(Account):
    """A bank account that charges for withdrawals and can deposit checks.

    >>> check = Check("Steven", 42)  # 42 dollars, payable to Steven
    >>> steven_account = CheckingAccount("Steven")
    >>> eric_account = CheckingAccount("Eric")
    >>> eric_account.deposit_check(check)  # trying to steal steven's money
    The police have been notified.
    >>> eric_account.balance
    0
    >>> steven_account.balance
    0
    >>> check.deposited
    False
    >>> eric_account.transactions
    []
    >>> steven_account.transactions
    []
    >>> steven_account.deposit_check(check)
    42
    >>> check.deposited
    True
    >>> steven_account.transactions
    [('deposit', 42)]
    >>> steven_account.deposit_check(check)  # can't cash check twice
    The police have been notified.
    >>> steven_account.transactions  # transactions history shouldn't have changed
    [('deposit', 42)]
    """
    withdraw_fee = 1
    interest = 0.01

    def withdraw(self, amount):
        "*** YOUR CODE HERE ***"

    def deposit_check(self, check):
        "*** YOUR CODE HERE ***"

```

```python
grader.check("checking_account")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Linked Lists

<hr style="border: 1px solid #fdb515;" />

A linked list is either an empty linked list (`Link.empty`) **or** a first value and the rest of the linked list.

```python
class Link:
    """
    >>> s = Link(1, Link(2, Link(3)))
    >>> s
    Link(1, Link(2, Link(3)))
    """
    empty = ()

    def __init__(self, first, rest=empty):
        assert rest is Link.empty or isinstance(rest, Link)
        self.first = first
        self.rest = rest

    def __repr__(self):
        if self.rest is not Link.empty:
            rest_str = ', ' + repr(self.rest)
        else:
            rest_str = ''
        return 'Link({0}{1})'.format(repr(self.first), rest_str)
```

To check if a `Link` is empty, compare it against the class attribute
`Link.empty`. For example, the below function prints out whether or not the link it is handed is empty:

```python
# Notice that this is not part of the Link class.
def test_empty(link):
    if link is Link.empty:
        print('This linked list is empty!')
    else:
        print('This linked list is not empty!')

# This function helps us print the elements of a linked list in a readable format.
def print_link(link):
    """Print elements of a linked list link.

    >>> link = Link(1, Link(2, Link(3)))
    >>> print_link(link)
    <1 2 3>
    >>> link1 = Link(1, Link(Link(2), Link(3)))
    >>> print_link(link1)
    <1 <2> 3>
    >>> link1 = Link(3, Link(Link(4), Link(5, Link(6))))
    >>> print_link(link1)
    <3 <4> 5 6>
    """
    print('<' + helper(link).rstrip() + '>')

def helper(link):
    if link == Link.empty:
        return ''
    elif isinstance(link.first, Link):
        return '<' + helper(link.first).rstrip() + '> ' + helper(link.rest)
    else:
        return str(link.first) +' '+  helper(link.rest)
```

```python
test_empty(Link(1, Link(2, Link(3))))  # This linked list is not empty!
```

```python
test_empty(Link.empty) # This linked list is empty!
```

```python
print_link(Link(1, Link(2, Link(3))))  # <1 2 3>
```

> Note: Linked lists are recursive data structures! A linked list contains the first element of the list (`first`) and a reference to another linked list (`rest`) which contains the rest of the values in the list.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Convert List

<hr style="border: 1px solid #fdb515;" />

Write a function `list_to_link` that takes in a Python list and returns the
linked list version of the sequence.

Try to find both an iterative and recursive solution for this problem!

```python
def list_to_link(lst):
    """Takes a Python list and returns a linked list with the same elements.

    >>> link = list_to_link([1, 2, 3])
    >>> print(link)
    <1 2 3>
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("list_to_link")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Reverse

<hr style="border: 1px solid #fdb515;" />

Implement `reverse`, which takes a linked list `link` and returns a linked list containing the elements of `link` in reverse order. The original `link` should be unchanged.

```python
def reverse(link):
    """Returns a Link that is the reverse of the original.

    >>> print_link(reverse(Link(1)))
    <1>
    >>> link = Link(1, Link(2, Link(3)))
    >>> new = reverse(link)
    >>> print_link(new)
    <3 2 1>
    >>> print_link(link)
    <1 2 3>
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("reverse")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Slice

<hr style="border: 1px solid #fdb515;" />

Implement a function `slice_link` that slices a given linked list `link`. `slice_link`
should slice the `link` starting at `start` and ending one element before
`end`, as with a normal Python list. Additionally, similar to slicing normal 
Python lists, this function should return a new list and not modify the original list.

```python
def slice_link(link, start, end):
    """Slices a linked list from start to end (as with a normal Python list).

    >>> link = Link(3, Link(1, Link(4, Link(1, Link(5, Link(9))))))
    >>> new = slice_link(link, 1, 4)
    >>> print(new)
    <1 4 1>
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("slice")
```

Tests for checking_account (points: 1)
```
>>> import io
>>> from contextlib import redirect_stdout
>>> check = Check('Steven', 42)
>>> steven_account = CheckingAccount('Steven')
>>> eric_account = CheckingAccount('Eric')
>>> with redirect_stdout(io.StringIO()):
...     eric_account.deposit_check(check)
>>> assert eric_account.balance == 0
>>> assert steven_account.balance == 0
>>> assert check.deposited is False
>>> assert eric_account.transactions == []
>>> assert steven_account.transactions == []
>>> with redirect_stdout(io.StringIO()):
...     result = steven_account.deposit_check(check)
>>> assert result == 42
>>> assert steven_account.balance == 42
>>> assert check.deposited is True
>>> assert steven_account.transactions == [('deposit', 42)]
>>> assert eric_account.transactions == []
>>> with redirect_stdout(io.StringIO()):
...     steven_account.deposit_check(check)
>>> assert steven_account.balance == 42
>>> assert steven_account.transactions == [('deposit', 42)]
```


Tests for list_to_link (points: 1)
```
>>> link = list_to_link([1, 2, 3])
>>> repr(link) == 'Link(1, Link(2, Link(3)))'
True
```


Tests for reverse (points: 1)
```
>>> print_link(reverse(Link(1)))
<1>
```
```
>>> link = Link(1, Link(2, Link(3)))
>>> new = reverse(link)
>>> (print_link(new), print_link(link))
<3 2 1>
<1 2 3>
(None, None)
```


Tests for slice (points: 1)
```
>>> link = Link(3, Link(1, Link(4, Link(1, Link(5, Link(9))))))
>>> new = slice_link(link, 1, 4)
>>> print_link(new)
<1 4 1>
```
