<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 2: Control

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# What Would Python Display? (WWPD): Control

<hr style="border: 1px solid #fdb515;" />

One component of lab assignments is to predict how the Python will behave. For the following questions, assign the variable to the **output** of various statements/expressions.


---

## 1)

```python
>>> def xk(c,d):
...     if c == 4:
...         return 6
...     elif d >= 4:
...         return 6 + 7 + c
...     else:
...         return 25
```

---

## 1a)

```python
>>> xk(10, 10)
```

```python
one_a = ...
```

```python
grader.check("q1a")
```

---

## 1b)

```python
>>> xk(10, 6)
```

```python
one_b = ...
```

```python
grader.check("q1b")
```

---

## 1c)

```python
>>> xk(4, 6)
```

```python
one_c = ...
```

```python
grader.check("q1c")
```

---

## 1d)

```python
>>> xk(0, 0)
```

```python
one_d = ...
```

```python
grader.check("q1d")
```

---

## 2)

```python
>>> def how_big(x):
...     if x > 10:
...         print('huge')
...     elif x > 5:
...         return 'big'
...     if x > 0:
...         print('positive')
...     else:
...         print(0)
```

---

## 2a)

```python
>>> how_big(7)
```

```python
two_a = ...
```

```python
grader.check("q2a")
```

---

## 2b)

```python
>>> print(how_big(7))
```

```python
two_b = ...
```

```python
grader.check("q2b")
```

---

## 2c and 2d)

```python
>>> how_big(12)
```

```python
two_c = ...
two_d = ...
```

```python
grader.check("q2cd")
```

---

## 2e - 2h)

```python
>>> print(how_big(12))
```

```python
two_e = ...
two_f = ...
two_g = ...
two_h = ...
```

```python
grader.check("q2eh")
```

---

## 2f - 2h)

``` python
print(how_big(1), how_big(0))
```

_**Note**_: To indicate two items in one line, do (item1, item2) with the parenthesis and comma.

```python
two_f = ...
two_g = ...
two_h2 = ...
```

```python
grader.check("q2fh")
```

---

# 2i - 2l)

```python 
>>> n = 3
>>> while n >= 0:
...     n -= 1
...     print(n)
```

```python
two_i = ...
two_j = ...
two_k = ...
two_l = ...
```

```python
grader.check("q2il")
```

---

# 2m - 2o)

```python
>>> negative = -12
>>> while negative:  # All numbers are true values except 0
...    if negative + 6:
...        print(negative)
...    negative += 3
```

```python
two_m = ...
two_n = ...
two_o = ...
```

```python
grader.check("q2m")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# WWPD: What If?

<hr style="border: 1px solid #fdb515;" />

---

# 3)

```python
>>> def ab(c, d):
...     if c > 5:
...         print(c)
...     elif c > 7:
...         print(d)
...     print('foo')
```

---

# 3a and 3b)

```python
>>> ab(10, 20)
```

```python
three_a = ...
three_b = ...
```

```python
grader.check("q3ab")
```

---

# 3c - 3e)

```python
>>> def bake(cake, make):
...     if cake == 0:
...         cake = cake + 1
...         print(cake)
...     if cake == 1:
...         print(make)
...     else:
...         return cake
...     return make
```

---

``` python
>>> bake(0, 29)
```

```python
three_c = ...
three_d = ...
three_e = ...
```

```python
grader.check("q3ce")
```

---

# 3f and 3g)

```python
>>> bake(1, "mashed potatoes")
```

```python
three_f = ...
three_g = ...
```

```python
grader.check("q3f")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Double Eights

<hr style="border: 1px solid #fdb515;" />

Write a function that takes in a number and determines if the digits contain two adjacent 8s.


```python
def double_eights(n):
    # Return true if n has two eights in a row.
    # >>> double_eights(8)
    # False
    # >>> double_eights(88)
    # True
    # >>> double_eights(2882)
    # True
    # >>> double_eights(880088)
    # True
    # >>> double_eights(12345)
    # False
    # >>> double_eights(80808080)
    # False
    "*** YOUR CODE HERE ***"

```

```python
grader.check("double_eights")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Falling Factorial

<hr style="border: 1px solid #fdb515;" />

Let's write a function `falling`, which is a "falling" factorial that takes two arguments, `n` and `k`, and returns the product of `k` consecutive numbers, starting from `n` and working downwards.

If `k` is larger than n, only multiply up to n consecutive numbers!

```python
def falling(n, k):
    # Compute the falling factorial of n to depth k.

    # >>> falling(6, 3)  # 6 * 5 * 4
    # 120
    # >>> falling(4, 3)  # 4 * 3 * 2
    # 24
    # >>> falling(4, 1)  # 4
    # 4
    # >>> falling(4, 0)
    # 1
    "*** YOUR CODE HERE ***"

```

```python
grader.check("falling_factorial")
```

Tests for double_eights (points: 1)
```
>>> double_eights(8) == False
True
```
```
>>> double_eights(88) == True
True
```
```
>>> double_eights(2882) == True
True
```
```
>>> double_eights(880088) == True
True
```
```
>>> double_eights(12345) == False
True
```
```
>>> double_eights(80808080) == False
True
```


Tests for falling_factorial (points: 1)
```
>>> falling(6, 3) == 120
True
```
```
>>> falling(4, 3) == 24
True
```
```
>>> falling(4, 1) == 4
True
```
```
>>> falling(4, 0) == 1
True
```


Tests for q1a (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_a).encode()).hexdigest()
'535fa30d7e25dd8a49f1536779734ec8286108d115da5045d77f3b4185d8f790'
```


Tests for q1b (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_b).encode()).hexdigest()
'535fa30d7e25dd8a49f1536779734ec8286108d115da5045d77f3b4185d8f790'
```


Tests for q1c (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_c).encode()).hexdigest()
'e7f6c011776e8db7cd330b54174fd76f7d0216b612387a5ffcfb81e6f0919683'
```


Tests for q1d (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_d).encode()).hexdigest()
'b7a56873cd771f2c446d369b649430b65a756ba278ff97ec81bb6f55b2e73569'
```


Tests for q2a (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(two_a.encode()).hexdigest()
'2a21fe6d592a19b7de898b50eb53c429608de1a66f3e9f62da19714a770553d1'
```


Tests for q2b (points: 1)
```
>>> two_b == print('big')
big
True
```


Tests for q2cd (points: 1)
```
>>> two_c == print('huge')
huge
True
```
```
>>> two_d == print('positive')
positive
True
```


Tests for q2eh (points: 1)
```
>>> two_e == print('huge')
huge
True
```
```
>>> two_f == print('positive')
positive
True
```
```
>>> two_g == None
True
```
```
>>> two_h == None
True
```


Tests for q2fh (points: 1)
```
>>> two_f == print('positive')
positive
True
```
```
>>> two_g == print('0')
0
True
```
```
>>> two_h2 == (None, None)
True
```


Tests for q2il (points: 1)
```
>>> two_i = print(2)
2
```
```
>>> two_j == print(1)
1
True
```
```
>>> two_k == print(0)
0
True
```
```
>>> two_f == print(-1)
-1
True
```


Tests for q2m (points: 1)
```
>>> two_m == print(-12)
-12
True
```
```
>>> two_n == print(-9)
-9
True
```
```
>>> two_o == print(-3)
-3
True
```


Tests for q3ab (points: 1)
```
>>> three_a == print('10')
10
True
```
```
>>> three_b == print('foo')
foo
True
```


Tests for q3ce (points: 1)
```
>>> three_c == print('1')
1
True
```
```
>>> three_d == print('29')
29
True
```
```
>>> three_e == 1
True
```


Tests for q3f (points: 1)
```
>>> three_f == print('mashed potatoes')
mashed potatoes
True
```
```
>>> three_g == 'mashed potatoes'
True
```
