# Lab 6

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS88 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

## Lab 6: Recursion and Tree Recursion

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## AB+C

<hr style="border: 1px solid #fdb515;" />

Implement `ab_plus_c`, a function that takes arguments `a`, `b`, and
`c` and computes `a * b + c`.  You can assume a and b are both non-negative
integers. However, you can't use the `*` operator. Use recursion!

```python
def ab_plus_c(a, b, c):
    """Computes a * b + c.
    >>> ab_plus_c(2, 4, 3)  # 2 * 4 + 3
    11
    >>> ab_plus_c(0, 3, 2)  # 0 * 3 + 2
    2
    >>> ab_plus_c(3, 0, 2)  # 3 * 0 + 2
    2
    """
    "*** YOUR CODE HERE ***"


```

```python
grader.check("ab_plus_c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Filter

<hr style="border: 1px solid #fdb515;" />

Write the recursive version of the function `filter` which returns a list and takes in

- `f` - a one-argument function that returns `True` if the passed in argument should be included in the resulting list or `False` otherwise
- `seq` - a list of values

Note that this is different from the built in `filter` function we learned previously, which returns a filter object, not a list.

```python
def is_even(x):
    return x % 2 == 0

def divisible_by5(x):
    return x % 5 == 0

def filter(f, seq):
    """Filter a sequence to only contain values allowed by filter.

    >>> def is_even(x):
    ...     return x % 2 == 0
    >>> def divisible_by5(x):
    ...     return x % 5 == 0
    >>> filter(is_even, [1,2,3,4])
    [2, 4]
    >>> filter(divisible_by5, [1, 4, 9, 16, 25, 100])
    [25, 100]
    >>> filter(is_even, [1])
    []
    >>> filter(is_even, [2])
    [2]
    >>> filter(is_even, [])
    []
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("filter")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Decimal

<hr style="border: 1px solid #fdb515;" />

Write the recursive version of the function `decimal` which takes in an integer `n` and returns a list of its digits, the decimal representation of `n`. See the doctests to handle the case where `n < 0`.

```python
def decimal(n):
    """Return a list representing the decimal representation of a number.

    >>> decimal(0)
    [0]
    >>> decimal(2)
    [2]
    >>> decimal(-8)
    ['-', 8]
    >>> decimal(-136)
    ['-', 1, 3, 6]
    >>> decimal(55055)
    [5, 5, 0, 5, 5]
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("decimal")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Paths

<hr style="border: 1px solid #fdb515;" />

Consider an insect in an *M* by *N* grid. The insect starts at the
bottom left corner, *(1, 1)*, and wants to end up at the top right
corner, *(M, N)*. The insect is only capable of moving right or
up. Write a function `paths` that takes a grid length and width
and returns the number of different paths the insect can take from the
start to the goal. (There is a [closed-form solution](https://en.wikipedia.org/wiki/Closed-form_expression) to this problem,
but try to answer it procedurally using recursion.)

<img src="grid.jpg" width="50%" alt="Grid diagrams illustrating sample paths from bottom-left to top-right using only right and up moves on 2x2 and 3x3 grids."/>

For example, the 2 by 2 grid has a total of two ways for the insect to
move from the start to the goal. For the 3 by 3 grid, the insect has 6
diferent paths (only 3 are shown above).

> **Hint:** What happens if we hit the top or rightmost edge?

```python
def paths(m, n):
    """Return the number of paths from one corner of an
    M by N grid to the opposite corner.

    >>> paths(2, 2)
    2
    >>> paths(5, 7)
    210
    >>> paths(117, 1)
    1
    >>> paths(1, 157)
    1
    """
    "*** YOUR CODE HERE ***"

```

```python
grader.check("paths")
```

Tests for ab_plus_c (points: 1)
```
>>> ab_plus_c(2, 4, 3) == 11
True
```
```
>>> ab_plus_c(0, 3, 2) == 2
True
```
```
>>> ab_plus_c(3, 0, 2) == 2
True
```


Tests for decimal (points: 1)
```
>>> assert decimal(0) == [0]
>>> assert decimal(2) == [2]
>>> assert decimal(-8) == ['-', 8]
>>> assert decimal(-136) == ['-', 1, 3, 6]
>>> assert decimal(55055) == [5, 5, 0, 5, 5]
```


Tests for filter (points: 1)
```
>>> assert filter(is_even, [1, 2, 3, 4]) == [2, 4]
>>> assert filter(divisible_by5, [1, 4, 9, 16, 25, 100]) == [25, 100]
>>> assert filter(is_even, [1]) == []
>>> assert filter(is_even, [2]) == [2]
>>> assert filter(is_even, []) == []
```


Tests for paths (points: 1)
```
>>> assert paths(2, 2) == 2
>>> assert paths(5, 7) == 210
>>> assert paths(117, 1) == 1
>>> assert paths(1, 157) == 1
```
